---
title: 'Deep Learning in TensorFlow: Increasing the Accuracy of your Models'
description: 'Chapter description goes here.'
---

## Lesson 4.1: Introducing Regularization Techniques

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Introducing Regularization Techniques

```yaml
type: VideoExercise
key: 6fb2d8d327
xp: 50
```

`@projector_key`
635cdc359ba2623063f2538725532e05
